```python
# Initialize Otter
import otter
grader = otter.Notebook("diy02.ipynb")
```

# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the "Run cell" button at the top that looks like ▶| or hold down `Shift` + `Return`.

# DIY02: Templates for Autograded and Manual Lessons

Welcome to the Summer Institute for Data Science! During this session will create notebooks like this one. You can't learn technical subjects without hands-on practice, so these activities are an important part of the session.

**Collaboration Policy:**

Collaborating on labs is more than okay $-$ it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too $-$ the best way to solidify your knowledge of a subject is to explain it. You should **_not just_** copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

**Due Date:** 

## Today's Assignment

In today's assignment, you'll learn how to:

- create a teacher notebook with autograded and manual questions

- create a student notebook from a teacher notebook.

Let's get started!

**Name:** 

**Course:** 

**Date:**

Run the cell below to set up the notebook.

```python
import numpy as np
from datascience import *

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

Let's load a few datasets. We'll use them to make up questions (autograded and manual) for our notebook.

```python
actors = Table.read_table('data/actors.csv')
hybrid = Table.read_table('data/hybrid.csv')
sons = Table.read_table('data/sons_heights.csv')
```

## 1. Autograded Questions

**Question 1.** Use the `actors` table to find the top 10 grossing movies in descending order. Save the results to a table named `top10_gross`.


```python
top10_gross = ...
top10_gross
```

```python
grader.check("q1")
```

**Question 2.** What is the average height of all the sons from the `sons` table. Save this value to the variable named `avg_son_height`.


```python
avg_son_height = np.mean(sons.column('son'))
avg_son_height
```

```python
grader.check("q2")
```

```python
hybrid.labels
```

**Question 3.** What is the _median_ msrp for the vehicles in the `hybrid` table. Save this value to the variable named `med_msrp`.


```python
med_msrp = np.median(hybrid.column('msrp'))
med_msrp
```

```python
grader.check("q3")
```

## 2. Manually Graded Questions

<!-- BEGIN QUESTION -->

**Question 4.** Have you seen any of the movies in the `top10_gross` table? If so, which one - which one is your favorite? If not, list a movie that you've seen or a book that you've reaf in the last 6 months. Explain why you like the movie or book.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.** Make a horizontal bar graph that shows the `Actor` and the `Gross` for the `top10_gross` table.

**Hint:** Replace the `____` with the correct method name and the ellipses (...) with the correct labels.

```python
top10_gross.___(...,...)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6.** Write a function that takes a sample (without replacement) of 50 cars from the `hybrid` table and returns the median msrp from the sample.

```python
def sample_med():
    """Takes a sample (without repalcement) of 50 cars 
        from the hybrid table and returns the median msrp 
        from the sample.
    """
    ...
    return ...
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> top10_gross != ...
True
```
failure_message: Did you forget to enter an expression?
```
>>> isinstance(top10_gross, tables.Table)
True
```
failure_message: Check your data type.
```
>>> top10_gross.num_rows == 10
True
```
failure_message: Do you have the correct number of rows?


Tests for q2 (points: None)
```
>>> avg_son_height != ...
True
```
failure_message: Did you forget to enter an expression?
```
>>> isinstance(avg_son_height, float)
True
```
failure_message: Check your data type.
```
>>> import math
>>> math.isclose(70.45, avg_son_height, rel_tol=0.001)
True
```
failure_message: Did you perform the correct calculation?
```
>>> avg_son_height
70.454748603351959
```
failure_message: Did you perform the correct calculation?


Tests for q3 (points: None)
```
>>> med_msrp != ...
True
```
failure_message: Did you forget to enter an expression?
```
>>> isinstance(med_msrp, float)
True
```
failure_message: Check your data type.
```
>>> med_msrp == 31950.0
True
```
failure_message: Check your calculation.
```
>>> med_msrp
31950.0
```
failure_message: Check your calculation.
