```python
# Initialize Otter
import otter
grader = otter.Notebook("diy00.ipynb")
```

# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the "Run cell" button at the top that looks like ▶| or hold down `Shift` + `Return`.

# DIY00: Example Question Formats

Welcome to the Summer Institute for Data Science! During this session will create notebooks like this one. You can't learn technical subjects without hands-on practice, so these activities are an important part of the session.

**Collaboration Policy:**

Collaborating on labs is more than okay $-$ it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too $-$ the best way to solidify your knowledge of a subject is to explain it. You should **_not just_** copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

**Due Date:** 

## Today's Assignment

In today's assignment, you'll learn how to:

- crate autograded questions

- create manually graded questions

- use the `otter assign` command to create a student notebook.

Let's get started!

**Name:** Mahmoud Harding

**Course:** Summer Institute

**Date:** 07/11/23

## 0. Assignment Config

The first cell of the notebook, places YAML-formatted configurations in a raw cell that begins with the comment `# ASSIGNMENT CONFIG`.

```
# ASSIGNMENT CONFIG
init_cell: false
export_cell: true
generate: true
files: 
 - data/hybrid.csv
 - images/hybrid-car.png
```

**Note:** Files that are required for the assignment must be listed in the assignment configuration block.


## 1. Autograded Question

Here are examples of an atuograded question:


**Question 1.** How many seconds are in a week?

```python
secs_in_week = 7 * 24 * 60 * 60 # SOLUTION
secs_in_week
```

Output:
```
604800
```

```python
grader.check("q1")
```

**Question 2.** How many seconds are in a week?

```python
secs_in_week = 7 * 24 * 60 * 60 # SOLUTION
secs_in_week
```

Output:
```
604800
```

```python
grader.check("q2")
```

**Question 3.** How many seconds are in a week?

```python
secs_in_week = 7 * 24 * 60 * 60 # SOLUTION
secs_in_week
```

Output:
```
604800
```

```python
grader.check("q3")
```

**Question 4.** How many seconds are in a week?

```python
secs_in_week = 7 * 24 * 60 * 60 # SOLUTION
```

```python
grader.check("q4")
```

Here is an example of what the student will see.

**Question 1.** How many seconds in a week?

`secs_in_week` = . . .

## 2. Manually Graded Question

Here is an example of a manually graded question:

<!-- BEGIN QUESTION -->

**Question 5.** What is your favorite day of the week?

_Type your answer here, replacing this text._

**SOLUTION:** My favorite day of the week is Friday. Answers may vary.

<!-- END QUESTION -->

Here is an example of what the student will see.

**Question 5.** What is your favorite day of the week?

_Type your answer here replacing this text._

## 3. Importing

Run the cell below to set up the notebook.

```python
from datascience import *
```

Read a dataset

```python
hybrid = Table.read_table("data/hybrid.csv")
```

```python
hybrid
```

Output:
```
vehicle         | year | msrp    | acceleration | mpg   | class
Prius (1st Gen) | 1997 | 24509.7 | 7.46         | 41.26 | Compact
Tino            | 2000 | 35355   | 8.2          | 54.1  | Compact
Prius (2nd Gen) | 2000 | 26832.2 | 7.97         | 45.23 | Compact
Insight         | 2000 | 18936.4 | 9.52         | 53    | Two Seater
Civic (1st Gen) | 2001 | 25833.4 | 7.04         | 47.04 | Compact
Insight         | 2001 | 19036.7 | 9.52         | 53    | Two Seater
Insight         | 2002 | 19137   | 9.71         | 53    | Two Seater
Alphard         | 2003 | 38084.8 | 8.33         | 40.46 | Minivan
Insight         | 2003 | 19137   | 9.52         | 53    | Two Seater
Civic           | 2003 | 14071.9 | 8.62         | 41    | Compact
... (143 rows omitted)
```

Format an image

```python

```

<img src="images/hybrid-car.png"/>

<img src="images/hybrid-car.png" width="200"/>

<img src="images/hybrid-car.png" height="800" width="1200"/>

<div style="text-align:center"><img src="images/hybrid-car.png" /></div>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> secs_in_week == 604800
True
```


Tests for q2 (points: None)
```
>>> secs_in_week == 604800
True
```
```
>>> type(secs_in_week) == int
True
```


Tests for q3 (points: None)
```
>>> secs_in_week != ...
True
```
```
>>> isinstance(secs_in_week, int)
True
```
```
>>> secs_in_week == 604800
True
```


Tests for q4 (points: None)
```
>>> secs_in_week != ...
True
```
failure_message: Did you forget to enter an expression?
```
>>> isinstance(secs_in_week, int)
True
```
failure_message: Check the data type of your result.
```
>>> secs_in_week == 604800
True
```
